Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **32 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Chretien Li"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Maximum subarray problem—part II

## Question 1 [time estimate: 2 minutes]
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
def bruteforce_max_subarray(A):
    #intialize all the values we're going to store as 0
    #this includes the total value (max_subarray), the initial index of subarray (index_i), final index (index_j), and for fun, the subarray itself
    max_subarray = float("-inf")
    index_i = 0
    index_j = 0
    max_index = []
    
    #nested for loop. What we're trying to do is for each value in our subarray, check the subarray for each index after it
    #so we start with the first value of array. Then check the value of the subarray formed with the number right after it, then the one after that and so on
    #if we come up with a higher value, store it
    #once we check for every value, repeat but starting at second value of array and so on
    for i in range(len(A)+1):
        
        #want to define j starting from next value after i, this way we don't return a subarray of just 1
        for j in range(i+1, len(A)+1):
            #temporary value of new subarray to check later if it's greater than what we have
            temp = 0
            
            #for loop to add up all values in subarray
            for x in A[i: j]:
                temp += x
                
            #check to see if it's larger than currently stored value. If so, store value alongside all other information we're interested in
            if temp > max_subarray:
                max_subarray = temp
                index_i = i
                index_j = j
                max_index = (tuple(A[i:j]))
    #return desired values. 
    return index_i, index_j-1, max_subarray

bruteforce_max_subarray([-2, -3, 4, -1, -2, 1])
```

Output:
```
(2, 2, 4)
```

## Question 2 [time estimate: 5 minutes]
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. With the output from ```max_subarray(A)```, do you need to rerun ```max_subarray()``` for the new array again to find the subarray we need? Explain why, or why not, in as much detail as possible, using, if necessary, a specific example to illustrate your reasoning. The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
Yes we do. Originally, for array A, the maximum subarray was just a single index, [4]. This is because all other values are negative and would lower the sum. The exception to this being the 1 at the end, but the negative values sandwhiched between "4" and "1" are greater in magnitude than just 1. However, now that 8 is added to subarray B, the maximum subarray is no longer just 4.Going from index [2,6], inclusive, the new maximum subarray, the sum is now 10, a larger value than previously in subarray A. THe algorithm as is has no way of determining this other than just doing everything single guess and check. As such, we have to run the algorithm again. If I understand the problem correctly, had the array been all negative values, then we wouldn't need to check everything again so long as the appended value is positive, as that automatically makes the new subarray (from 4 until the end) the maxiumum sub array.

## Question 3 [time estimate: 5 minutes]

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
False. Without 8, the subarray value remains at 4. Only with 8 included can we get the new value of 10

## Question 4 [time estimate: 10 minutes]
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(x, mx):
    """ 
    Parameters
    ----------
    x : list
        A NON-EMPTY list of numbers (e.g., the 
        array B in the first two questions above)
        
        If x has 1 element: returns the value of 
        the element regardless of the value of mx
        
    mx 
        The maximum subarray of x excluding its 
        ast element (i.e., compute the maximum 
        subarray of the input array x considering 
        only its first len(x) - 1 elements)
    
    Returns
    -------
    int
        The maximum subarray of the array x
    
    Example
    -------
    Using the array B in question 2, the result 
    of incremental_max_subarray(B, 4) is 10: 
    
        10 = 8 + 1 - 2 -1 + 4 
    
    
    """
    # if our input x only has one input, that number is by default the maximum value
    if len(x) == 1:
        return x[0]
    else:
        #initialize a value that serves as a tracker. it will be used to hold the maximum value of the array
        #Baring that this value remains above zero, if it doesn't, then reset tracker value to 0 as any subarray including slice up until this point won't have positive contribution
        count = 0
        
        #another value stores the maximum value so long as it was larger than the previous one. This way, if the value dips below zero, never to come back up again, we know what the largest value was
        max_subarray = float("-inf")
        
        #iterate over every item in inputted list
        for i in x:
            
            #add item to count
            count += i
            
            #if the count at this point is larger than previous maximum value, set count as new largest value
            if count > max_subarray:
                max_subarray = count
            #if count goes below zero, reset to 0
            if count < 0:
                count = 0
        
        #return the larger of maximum value without last entry, and maximum
        return max(max_subarray, mx)
    
```

```python
B = [-2, -3, 4, -1, -2, 1, 8]
print(B[:1])
```

Output:
```
[-2]
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

## Question 5

### Part 1 [time estimate: 5 minutes]

Now use `incremental_max_subarray(x, mx)` iteratively, starting from the first element, to calculate the maximum subarray of a given list. Come up with a few test cases to test your code.

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    #same as above, but only return overall maximum value of inputed list
    # if our input x only has one input, that number is by default the maximum value
    if len(A) == 1:
        return x[0]
    else:
        #initialize a value that serves as a tracker. it will be used to hold the maximum value of the array
        #Baring that this value remains above zero, if it doesn't, then reset tracker value to 0 as any subarray including slice up until this point won't have positive contribution
        count = 0
        
        #another value stores the maximum value so long as it was larger than the previous one. This way, if the value dips below zero, never to come back up again, we know what the largest value was
        max_subarray = float("-inf")
        
        #iterate over every item in inputted list
        for i in A:
            
            #add item to count
            count += i
            
            #if the count at this point is larger than previous maximum value, set count as new largest value
            if count > max_subarray:
                max_subarray = count
            #if count goes below zero, reset to 0
            if count < 0:
                count = 0
        
        #return the larger of maximum value without last entry, and maximum
        return max_subarray
    
print (max_subarray([-2, -3, 4, -1, -2, 1]))
```

Output:
```
4
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2 [time estimate: 5 minutes]
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
Yes this is more efficient that divide and conquer. In terms of time complexity, divide and conquer has a complexity of O(N) + O(Nlog(N)), while this approach O(N). As all we need to do is iterate across each value of the list in a linear fashion, this approach is more efficient